<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Q3 Medium B — Laplace Transform Application to First-Order   -->
<!-- Circuits. 4 variants, 13 pts each, ~30 min.                  -->
<!-- Complex resistor networks (3-4 nodes, 3 meshes).             -->
<!-- Subparts: A (2 pts MCQ), B (4 pts Numerical Rth/Req),        -->
<!--           C (4 pts Numerical tau), D (3 pts Essay)           -->
<!-- ============================================================ -->

<!-- ============================================================ -->
<!-- VARIANT 1: RL Step — Pi-network feeding inductor              -->
<!-- Vs→R1→(A), R2: A→GND, R3: A→(B), L: B→GND                  -->
<!-- Rth = R3 + R1||R2, tau = L/Rth                               -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V1_RL_Step_Network</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Question 3 — First-Order Circuit Analysis Using Laplace Transforms (13 pts)</h3>

<p>Consider the circuit shown below. At \(t = 0\), a DC voltage source \(V_s = {@Vs@}\;\text{V}\) is applied as a step input \(V_s \cdot u(t)\). The circuit contains resistors \(R_1 = {@R1@}\;\Omega\), \(R_2 = {@R2@}\;\Omega\), \(R_3 = {@R3@}\;\Omega\), and an inductor \(L = {@L@}\;\text{H}\). Assume <strong>zero initial conditions</strong>.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v1_rl_step.svg]</strong></p>

<p><em><strong>Note:</strong> Express numerical answers to 3 significant figures. Scientific notation is accepted (e.g., <code>2.50*10^3</code> or <code>2.50E3</code>).</em></p>

<hr/>

<h4>Part A (2 pts)</h4>
<p>When transforming this circuit to the s-domain with zero initial conditions, what is the s-domain impedance of the inductor \(L\)?</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <ul>
        <li>(a) \(sL\)</li>
        <li>(b) \(\dfrac{1}{sL}\)</li>
        <li>(c) \(\dfrac{L}{s}\)</li>
        <li>(d) \(s^2 L\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts)</h4>
<p>Determine the equivalent resistance \(R_\text{eq}\) that the inductor \(L\) "sees" looking back into the rest of the circuit from its terminals (with the voltage source replaced by its internal impedance).</p>
<p>Give your answer in ohms to 3 significant figures.</p>
<p>\(R_\text{eq} = \) [[input:ansB]] \(\Omega\) [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts)</h4>
<p>Using your result from Part B, compute the time constant \(\tau\) of this circuit.</p>
<p>Give your answer in seconds to 3 significant figures.</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Physical Interpretation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>What is the inductor current at \(t = 0^+\)? What is it as \(t \to \infty\)? Justify both using circuit analysis or the s-domain.</li>
    <li>Explain the physical meaning of the time constant \(\tau\) you computed. What happens to the current at \(t = \tau\)?</li>
</ul>
<p>[[input:ansD]] [[validation:ansD]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution</h4>
<p><strong>Part A:</strong> With zero initial conditions, an inductor \(L\) in the s-domain is represented by impedance \(Z_L = sL\). Answer: <strong>(a)</strong>.</p>
<p><strong>Part B:</strong> To find the equivalent resistance seen by L, replace Vs with a short circuit. Looking from the L terminals (node B to ground): R3 is in series with the parallel combination R1‖R2.
\[R_\text{eq} = R_3 + \frac{R_1 R_2}{R_1 + R_2} = {@R3@} + \frac{{@R1@} \cdot {@R2@}}{{@R1@} + {@R2@}} = {@Req@}\;\Omega\]</p>
<p><strong>Part C:</strong> \(\tau = L / R_\text{eq} = {@L@} / {@Req@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>At \(t=0^+\), \(i_L = 0\) because an inductor opposes instantaneous changes in current (zero initial conditions).</li>
    <li>As \(t\to\infty\), the inductor acts as a short circuit. The Thevenin voltage at node B is \(V_\text{th} = V_s \cdot R_2/(R_1+R_2) = {@Vth@}\;\text{V}\). The steady-state current is \(i_L(\infty) = V_\text{th}/R_\text{eq} = {@Isteady@}\;\text{A}\).</li>
    <li>The time constant \(\tau = {@tau@}\;\text{s}\) is the time at which the current reaches approximately 63.2% of its final value. After \(5\tau\), the current is within 1% of steady state.</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* RL Step with pi-network */
Vs: rand([8, 10, 12, 15]);
R1: rand([200, 330, 470, 680]);
R2: rand([100, 220, 470, 1000]);
R3: rand([100, 150, 220, 330]);
L: rand([0.02, 0.05, 0.1, 0.2]);

/* Computed quantities */
R1parR2: R1*R2/(R1+R2);
Req: R3 + R1parR2;
tau: L/Req;
Vth: Vs*R2/(R1+R2);
Isteady: Vth/Req;

/* Answers */
correctA: a;
correctB: Req;
correctC: tau;
taD: "See model answer in general feedback.";
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V1 RL Step: Vs={@Vs@}, R1={@R1@}, R2={@R2@}, R3={@R3@}, L={@L@}, Req={@Req@}, tau={@tau@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RL step response with resistor network. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d,A,B,C,D</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>numerical</type>
        <tans>correctB</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansD</name>
        <type>essay</type>
        <tans>taD</tans>
        <boxsize>15</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>0</mustverify>
        <showvalidation>0</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. Review the s-domain impedance models for inductors.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_ratio: if (correctB = 0 or ansB = 0) then 0 else float(abs(ansB / correctB));
_log10r: if _ratio > 0 then float(log(_ratio)/log(10)) else -999;
_nearint: round(_log10r);
_mag_err: is(_ratio > 0 and abs(_nearint) >= 1 and abs(_log10r - _nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtB-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Think carefully about which resistors are in series and which are in parallel when viewed from the inductor terminals.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_c_ratio: if (correctC = 0 or ansC = 0) then 0 else float(abs(ansC / correctC));
_c_log10r: if _c_ratio > 0 then float(log(_c_ratio)/log(10)) else -999;
_c_nearint: round(_c_log10r);
_c_mag_err: is(_c_ratio > 0 and abs(_c_nearint) >= 1 and abs(_c_log10r - _c_nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtC-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_c_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Recall the relationship between the time constant, the equivalent resistance, and the energy storage element.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <deployedseed>12345</deployedseed>
    <deployedseed>23456</deployedseed>
    <deployedseed>34567</deployedseed>
    <deployedseed>45678</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 2: RC Step — C bridges two resistor dividers          -->
<!-- Vs→R1→(A), R2: A→GND, C: A→(B), R3: B→GND(via right),      -->
<!-- R4: B→GND. Rth = (R1||R2) + (R3||R4), tau = Rth*C           -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V2_RC_Step_Bridge</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Question 3 — First-Order Circuit Analysis Using Laplace Transforms (13 pts)</h3>

<p>Consider the circuit shown below. At \(t = 0\), a DC voltage source \(V_s = {@Vs@}\;\text{V}\) is applied as a step input \(V_s \cdot u(t)\). The circuit contains resistors \(R_1 = {@R1@}\;\Omega\), \(R_2 = {@R2@}\;\Omega\), \(R_3 = {@R3@}\;\Omega\), \(R_4 = {@R4@}\;\Omega\), and a capacitor \(C = {@Cval@}\;\text{F}\). Assume <strong>zero initial conditions</strong>.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v2_rc_step.svg]</strong></p>

<p><em><strong>Note:</strong> Express numerical answers to 3 significant figures. Scientific notation is accepted (e.g., <code>2.50*10^3</code> or <code>2.50E3</code>).</em></p>

<hr/>

<h4>Part A (2 pts)</h4>
<p>When transforming this circuit to the s-domain with zero initial conditions, what is the s-domain impedance of the capacitor \(C\)?</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <ul>
        <li>(a) \(sC\)</li>
        <li>(b) \(\dfrac{1}{sC}\)</li>
        <li>(c) \(\dfrac{C}{s}\)</li>
        <li>(d) \(\dfrac{s}{C}\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts)</h4>
<p>Determine the equivalent resistance \(R_\text{eq}\) that the capacitor \(C\) "sees" looking out from its terminals into the rest of the circuit (with the voltage source replaced by its internal impedance).</p>
<p>Give your answer in ohms to 3 significant figures.</p>
<p>\(R_\text{eq} = \) [[input:ansB]] \(\Omega\) [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts)</h4>
<p>Using your result from Part B, compute the time constant \(\tau\) of this circuit.</p>
<p>Give your answer in seconds to 3 significant figures.</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Physical Interpretation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>What is the voltage across the capacitor at \(t = 0^+\)? What is it as \(t \to \infty\)? Justify both answers.</li>
    <li>Explain the physical meaning of the time constant \(\tau\) you computed. At \(t = \tau\), what fraction of the final voltage has the capacitor reached?</li>
</ul>
<p>[[input:ansD]] [[validation:ansD]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution</h4>
<p><strong>Part A:</strong> With zero initial conditions, a capacitor \(C\) in the s-domain has impedance \(Z_C = 1/(sC)\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> Remove C and short the voltage source. From the A-side terminal: R1‖R2. From the B-side terminal: R3‖R4. Since these are in series:
\[R_\text{eq} = \frac{R_1 R_2}{R_1+R_2} + \frac{R_3 R_4}{R_3+R_4} = {@Req@}\;\Omega\]</p>
<p><strong>Part C:</strong> \(\tau = R_\text{eq} \cdot C = {@Req@} \times {@Cval@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>At \(t=0^+\), \(v_C = 0\) because a capacitor opposes instantaneous changes in voltage (zero initial conditions).</li>
    <li>As \(t\to\infty\), the capacitor is fully charged and no current flows through it. The Thevenin voltage is \(V_\text{th} = V_s \cdot R_2/(R_1+R_2) = {@Vth@}\;\text{V}\) (only the A-side divider contributes since no current flows through R3/R4 at steady state).</li>
    <li>At \(t = \tau\), the capacitor has reached approximately 63.2% of its final voltage (\(1-e^{-1}\)).</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* RC Step with bridging capacitor */
Vs: rand([5, 10, 12, 15]);
R1: rand([1000, 2200, 3300, 4700]);
R2: rand([1000, 2200, 4700, 6800]);
R3: rand([470, 1000, 1500, 2200]);
R4: rand([680, 1000, 1500, 3300]);
Cval: rand([10*10^(-6), 22*10^(-6), 47*10^(-6), 100*10^(-6)]);

/* Computed quantities */
R1parR2: R1*R2/(R1+R2);
R3parR4: R3*R4/(R3+R4);
Req: R1parR2 + R3parR4;
tau: Req*Cval;
Vth: Vs*R2/(R1+R2);

/* Answers */
correctA: b;
correctB: Req;
correctC: tau;
taD: "See model answer in general feedback.";
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V2 RC Step Bridge: Vs={@Vs@}, R1={@R1@}, R2={@R2@}, R3={@R3@}, R4={@R4@}, C={@Cval@}, Req={@Req@}, tau={@tau@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RC step response with bridging capacitor. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d,A,B,C,D</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>numerical</type>
        <tans>correctB</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansD</name>
        <type>essay</type>
        <tans>taD</tans>
        <boxsize>15</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>0</mustverify>
        <showvalidation>0</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. Review the s-domain impedance models for capacitors.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_ratio: if (correctB = 0 or ansB = 0) then 0 else float(abs(ansB / correctB));
_log10r: if _ratio > 0 then float(log(_ratio)/log(10)) else -999;
_nearint: round(_log10r);
_mag_err: is(_ratio > 0 and abs(_nearint) >= 1 and abs(_log10r - _nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtB-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Consider what the capacitor sees from each of its terminals independently.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_c_ratio: if (correctC = 0 or ansC = 0) then 0 else float(abs(ansC / correctC));
_c_log10r: if _c_ratio > 0 then float(log(_c_ratio)/log(10)) else -999;
_c_nearint: round(_c_log10r);
_c_mag_err: is(_c_ratio > 0 and abs(_c_nearint) >= 1 and abs(_c_log10r - _c_nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtC-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_c_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Recall the relationship between the time constant, the equivalent resistance, and the energy storage element.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <deployedseed>56789</deployedseed>
    <deployedseed>67890</deployedseed>
    <deployedseed>78901</deployedseed>
    <deployedseed>89012</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 3: RL Natural — Inductor with branching discharge     -->
<!-- L(I0) at X, R4: X→GND, R1: X→(A), R2: A→GND, R3: A→GND    -->
<!-- Req = R4 || [R1 + R2||R3], tau = L/Req                      -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V3_RL_Natural_Network</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Question 3 — First-Order Circuit Analysis Using Laplace Transforms (13 pts)</h3>

<p>In the circuit shown below, an inductor \(L = {@L@}\;\text{H}\) carries an initial current \(I_0 = {@I0@}\;\text{A}\) at \(t = 0\). At \(t = 0\), the inductor is connected to the resistor network consisting of \(R_1 = {@R1@}\;\Omega\), \(R_2 = {@R2@}\;\Omega\), \(R_3 = {@R3@}\;\Omega\), and \(R_4 = {@R4@}\;\Omega\). There is <strong>no external source</strong>.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v3_rl_natural.svg]</strong></p>

<p><em><strong>Note:</strong> Express numerical answers to 3 significant figures. Scientific notation is accepted (e.g., <code>2.50*10^3</code> or <code>2.50E3</code>).</em></p>

<hr/>

<h4>Part A (2 pts)</h4>
<p>In the s-domain, how is an inductor \(L\) with a nonzero initial current \(I_0\) modeled?</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <ul>
        <li>(a) Impedance \(sL\) only</li>
        <li>(b) Impedance \(sL\) in series with a voltage source \(LI_0\)</li>
        <li>(c) Impedance \(\dfrac{1}{sL}\) in parallel with a current source \(I_0/s\)</li>
        <li>(d) Impedance \(sL\) in series with a voltage source \(I_0/s\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts)</h4>
<p>Determine the equivalent resistance \(R_\text{eq}\) that the inductor "sees" looking into the resistor network from its terminals.</p>
<p>Give your answer in ohms to 3 significant figures.</p>
<p>\(R_\text{eq} = \) [[input:ansB]] \(\Omega\) [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts)</h4>
<p>Using your result from Part B, compute the time constant \(\tau\) of this circuit.</p>
<p>Give your answer in seconds to 3 significant figures.</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Physical Interpretation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>What is the inductor current at \(t = 0^+\)? What is it as \(t \to \infty\)? Justify both answers.</li>
    <li>The initial energy stored in the inductor is \(W_0 = \frac{1}{2}LI_0^2\). Where does this energy go? Explain how it is distributed among the resistors in the network.</li>
</ul>
<p>[[input:ansD]] [[validation:ansD]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution</h4>
<p><strong>Part A:</strong> An inductor with initial current \(I_0\) is modeled as impedance \(sL\) in series with a voltage source \(LI_0\). This arises from \(\mathcal{L}\{L\,di/dt\} = sL \cdot I(s) - LI_0\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> Looking from the inductor terminals (node X to ground): R4 is directly from X to ground. In parallel with that path, R1 goes from X to node A, where R2 and R3 are both connected to ground (R2‖R3). So:
\[R_\text{eq} = R_4 \;\|\; \left[R_1 + (R_2 \| R_3)\right] = R_4 \;\|\; \left[R_1 + \frac{R_2 R_3}{R_2+R_3}\right] = {@Req@}\;\Omega\]</p>
<p><strong>Part C:</strong> \(\tau = L / R_\text{eq} = {@L@} / {@Req@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>At \(t=0^+\), \(i_L = I_0 = {@I0@}\;\text{A}\) because an inductor opposes instantaneous changes in current.</li>
    <li>As \(t\to\infty\), \(i_L = 0\) because there is no source to sustain the current.</li>
    <li>The stored energy \(W_0 = \frac{1}{2}({@L@})({@I0@})^2 = {@W0@}\;\text{J}\) is entirely dissipated as heat in the four resistors. The energy splits among R1, R2, R3, and R4 according to the current distribution through each path.</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* RL Natural with branching discharge */
I0: rand([0.5, 1, 2, 3]);
R1: rand([100, 220, 330, 470]);
R2: rand([220, 470, 680, 1000]);
R3: rand([150, 330, 560, 820]);
R4: rand([100, 270, 470, 680]);
L: rand([0.01, 0.02, 0.05, 0.1]);

/* Computed quantities */
R2parR3: R2*R3/(R2+R3);
R1plusR2parR3: R1 + R2parR3;
Req: R4*R1plusR2parR3/(R4+R1plusR2parR3);
tau: L/Req;
W0: (1/2)*L*I0^2;

/* Answers */
correctA: b;
correctB: Req;
correctC: tau;
taD: "See model answer in general feedback.";
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V3 RL Natural Network: I0={@I0@}, R1={@R1@}, R2={@R2@}, R3={@R3@}, R4={@R4@}, L={@L@}, Req={@Req@}, tau={@tau@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RL natural response with branching resistor network. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d,A,B,C,D</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>numerical</type>
        <tans>correctB</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansD</name>
        <type>essay</type>
        <tans>taD</tans>
        <boxsize>15</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>0</mustverify>
        <showvalidation>0</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. Review the s-domain model for an inductor with nonzero initial current.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_ratio: if (correctB = 0 or ansB = 0) then 0 else float(abs(ansB / correctB));
_log10r: if _ratio > 0 then float(log(_ratio)/log(10)) else -999;
_nearint: round(_log10r);
_mag_err: is(_ratio > 0 and abs(_nearint) >= 1 and abs(_log10r - _nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtB-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Identify the different current paths from the inductor terminals to ground and combine them correctly.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_c_ratio: if (correctC = 0 or ansC = 0) then 0 else float(abs(ansC / correctC));
_c_log10r: if _c_ratio > 0 then float(log(_c_ratio)/log(10)) else -999;
_c_nearint: round(_c_log10r);
_c_mag_err: is(_c_ratio > 0 and abs(_c_nearint) >= 1 and abs(_c_log10r - _c_nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtC-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_c_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Recall the relationship between the time constant, the equivalent resistance, and the energy storage element.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <deployedseed>11111</deployedseed>
    <deployedseed>22222</deployedseed>
    <deployedseed>33333</deployedseed>
    <deployedseed>44444</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 4: RC Natural — T-network discharge                   -->
<!-- C(V0) at X, R1: X→(A), R3: A→GND, R2: A→(B), R4: B→GND    -->
<!-- Req = R1 + R3||(R2+R4), tau = Req*C                         -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V4_RC_Natural_Network</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Question 3 — First-Order Circuit Analysis Using Laplace Transforms (13 pts)</h3>

<p>In the circuit shown below, a capacitor \(C = {@Cval@}\;\text{F}\) is initially charged to a voltage \(V_0 = {@V0@}\;\text{V}\). At \(t = 0\), the capacitor is connected to the resistor network consisting of \(R_1 = {@R1@}\;\Omega\), \(R_2 = {@R2@}\;\Omega\), \(R_3 = {@R3@}\;\Omega\), and \(R_4 = {@R4@}\;\Omega\). There is <strong>no external source</strong>.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v4_rc_natural.svg]</strong></p>

<p><em><strong>Note:</strong> Express numerical answers to 3 significant figures. Scientific notation is accepted (e.g., <code>2.50*10^3</code> or <code>2.50E3</code>).</em></p>

<hr/>

<h4>Part A (2 pts)</h4>
<p>In the s-domain, how is a capacitor \(C\) with an initial voltage \(V_0\) modeled?</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <ul>
        <li>(a) Impedance \(\dfrac{1}{sC}\) only</li>
        <li>(b) Impedance \(\dfrac{1}{sC}\) in series with a voltage source \(\dfrac{V_0}{s}\)</li>
        <li>(c) Impedance \(sC\) in series with a voltage source \(V_0 \cdot s\)</li>
        <li>(d) Impedance \(\dfrac{1}{sC}\) in parallel with a current source \(\dfrac{V_0}{R_1}\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts)</h4>
<p>Determine the equivalent resistance \(R_\text{eq}\) that the capacitor "sees" looking into the resistor network from its terminals.</p>
<p>Give your answer in ohms to 3 significant figures.</p>
<p>\(R_\text{eq} = \) [[input:ansB]] \(\Omega\) [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts)</h4>
<p>Using your result from Part B, compute the time constant \(\tau\) of this circuit.</p>
<p>Give your answer in seconds to 3 significant figures.</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Physical Interpretation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>What is the capacitor voltage at \(t = 0^+\)? What is it as \(t \to \infty\)? Justify both answers.</li>
    <li>The initial energy stored in the capacitor is \(W_0 = \frac{1}{2}CV_0^2\). Where does this energy go? Explain how it is distributed among the resistors in the network.</li>
</ul>
<p>[[input:ansD]] [[validation:ansD]]</p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution</h4>
<p><strong>Part A:</strong> A capacitor with initial voltage \(V_0\) is modeled as impedance \(1/(sC)\) in series with a voltage source \(V_0/s\). This arises from \(V_C(s) = \frac{1}{sC}I(s) + \frac{V_0}{s}\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> Looking from the capacitor terminals (node X to ground): R1 goes from X to node A. From node A, there are two paths to ground: R3 directly, and R2+R4 in series through node B. So:
\[R_\text{eq} = R_1 + R_3 \;\|\; (R_2 + R_4) = R_1 + \frac{R_3(R_2+R_4)}{R_3+R_2+R_4} = {@Req@}\;\Omega\]</p>
<p><strong>Part C:</strong> \(\tau = R_\text{eq} \cdot C = {@Req@} \times {@Cval@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>At \(t=0^+\), \(v_C = V_0 = {@V0@}\;\text{V}\) because a capacitor opposes instantaneous changes in voltage.</li>
    <li>As \(t\to\infty\), \(v_C = 0\) because there is no source to maintain the voltage and all energy is dissipated.</li>
    <li>The stored energy \(W_0 = \frac{1}{2}({@Cval@})({@V0@})^2 = {@W0@}\;\text{J}\) is entirely dissipated as heat in the four resistors. The distribution depends on the current paths: more energy goes to resistors carrying more current.</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* RC Natural with T-network discharge */
V0: rand([5, 10, 12, 15]);
R1: rand([470, 1000, 1500, 2200]);
R2: rand([330, 680, 1000, 1500]);
R3: rand([1000, 2200, 3300, 4700]);
R4: rand([470, 1000, 1500, 2200]);
Cval: rand([10*10^(-6), 22*10^(-6), 47*10^(-6), 100*10^(-6)]);

/* Computed quantities */
R2plusR4: R2 + R4;
R3parR2R4: R3*R2plusR4/(R3+R2plusR4);
Req: R1 + R3parR2R4;
tau: Req*Cval;
W0: (1/2)*Cval*V0^2;

/* Answers */
correctA: b;
correctB: Req;
correctC: tau;
taD: "See model answer in general feedback.";
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V4 RC Natural Network: V0={@V0@}, R1={@R1@}, R2={@R2@}, R3={@R3@}, R4={@R4@}, C={@Cval@}, Req={@Req@}, tau={@tau@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RC natural response with T-network discharge. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d,A,B,C,D</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>numerical</type>
        <tans>correctB</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansD</name>
        <type>essay</type>
        <tans>taD</tans>
        <boxsize>15</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>0</mustverify>
        <showvalidation>0</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. Review the s-domain model for a capacitor with nonzero initial voltage.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_ratio: if (correctB = 0 or ansB = 0) then 0 else float(abs(ansB / correctB));
_log10r: if _ratio > 0 then float(log(_ratio)/log(10)) else -999;
_nearint: round(_log10r);
_mag_err: is(_ratio > 0 and abs(_nearint) >= 1 and abs(_log10r - _nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtB-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Trace the paths from the capacitor terminals through the resistor network to ground and combine them correctly.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text><![CDATA[
/* Order-of-magnitude check for tiered grading */
_c_ratio: if (correctC = 0 or ansC = 0) then 0 else float(abs(ansC / correctC));
_c_log10r: if _c_ratio > 0 then float(log(_c_ratio)/log(10)) else -999;
_c_nearint: round(_c_log10r);
_c_mag_err: is(_c_ratio > 0 and abs(_c_nearint) >= 1 and abs(_c_log10r - _c_nearint) < 0.2);
]]></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check within 5% (full marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.05</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Correct.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>1</name>
            <description>Check within 15% (70% marks)</description>
            <answertest>NumRelative</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.15</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.7</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-2-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your answer is close but not precise enough. You received 70% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>2</falsenextnode>
            <falseanswernote>prtC-2-F</falseanswernote>
            <falsefeedback format="html"><text></text></falsefeedback>
        </node>
        <node>
            <name>2</name>
            <description>Check order of magnitude error (60% marks)</description>
            <answertest>AlgEquiv</answertest>
            <sans>_c_mag_err</sans>
            <tans>true</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>0.6</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-3-T</trueanswernote>
            <truefeedback format="html"><text><![CDATA[<p>Your significant figures appear correct but the order of magnitude is wrong. Check your unit conversions. You received 60% for this part.</p>]]></text></truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-3-F</falseanswernote>
            <falsefeedback format="html"><text><![CDATA[<p>Incorrect. Recall the relationship between the time constant, the equivalent resistance, and the energy storage element.</p>]]></text></falsefeedback>
        </node>
    </prt>
    <deployedseed>55555</deployedseed>
    <deployedseed>66666</deployedseed>
    <deployedseed>77777</deployedseed>
    <deployedseed>88888</deployedseed>
</question>

</quiz>
