<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Q3 Medium B — Laplace Transform Application to First-Order   -->
<!-- Circuits. 4 variants, 13 pts each, ~30 min.                  -->
<!-- Subparts: A (2 pts MCQ), B (4 pts Algebraic/Numerical),      -->
<!--           C (4 pts Numerical), D (3 pts Essay)               -->
<!-- ============================================================ -->

<!-- ============================================================ -->
<!-- VARIANT 1: RL Step Response                                   -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V1_RL_Step_Response</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Q3 — Laplace Transform: RL Step-Response Circuit (13 pts)</h3>

<p>A series RL circuit consists of a resistor \(R = {@R@}\;\Omega\) and an inductor \(L = {@L@}\;\text{H}\) connected in series. At \(t = 0\), a DC voltage source \(V_s = {@Vs@}\;\text{V}\) is applied as a step input \(V_s \cdot u(t)\). Assume <strong>zero initial conditions</strong> (no current flowing before \(t=0\)).</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v1_rl_step.svg]</strong></p>

<hr/>

<h4>Part A (2 pts) — S-Domain Component Model (MCQ)</h4>
<p>When transforming the circuit to the s-domain with zero initial conditions, how is the inductor \(L\) modeled?</p>
<p>Select the correct s-domain impedance of the inductor:</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <p><strong>Options:</strong></p>
    <ul>
        <li>(a) \(sL\)</li>
        <li>(b) \(\dfrac{1}{sL}\)</li>
        <li>(c) \(\dfrac{L}{s}\)</li>
        <li>(d) \(s^2 L\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts) — Transfer Function</h4>
<p>Using the s-domain model, the current in the circuit is \(I(s)\) and the input voltage transform is \(V(s)\).</p>
<p>Determine the transfer function \(H(s) = \dfrac{I(s)}{V(s)}\).</p>
<p>Express your answer as a rational function of \(s\). Use the variable <code>s</code>.</p>
<p>For example, if your answer were \(\dfrac{3}{s+5}\), type <code>3/(s+5)</code>.</p>
<p>\(H(s) = \) [[input:ansB]] [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts) — Time Constant</h4>
<p>Compute the time constant \(\tau\) of this RL circuit.</p>
<p>Give your answer in seconds as a decimal number (e.g., <code>0.0005</code>).</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Initial/Final Value Check &amp; Physical Meaning (Essay)</h4>
<p>Using the <strong>Initial Value Theorem</strong> and the <strong>Final Value Theorem</strong>, verify the behaviour of the current \(i(t)\):</p>
<ul>
    <li>Show that \(i(0^+) = 0\) using the Initial Value Theorem on \(I(s)\).</li>
    <li>Show that \(i(\infty) = V_s/R = {@Vs/R@}\;\text{A}\) using the Final Value Theorem.</li>
    <li>Explain the <strong>physical meaning</strong> of the time constant \(\tau = {@tau@}\;\text{s}\) in this circuit. What fraction of the final current is reached at \(t = \tau\)?</li>
</ul>
<p><strong>Write your answer below:</strong></p>
<p><textarea name="essayD_v1" rows="12" cols="80" style="width:100%; font-family:monospace;">
</textarea></p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution — Variant 1 (RL Step Response)</h4>
<p><strong>Part A:</strong> With zero initial conditions, an inductor \(L\) in the s-domain is represented by impedance \(Z_L = sL\). Answer: <strong>(a)</strong>.</p>
<p><strong>Part B:</strong> KVL in s-domain: \(V(s) = I(s)(R + sL)\), so
\[H(s) = \frac{I(s)}{V(s)} = \frac{1}{R + sL} = \frac{1}{{@R@} + s \cdot {@L@}}\]</p>
<p><strong>Part C:</strong> The time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>Initial Value Theorem: \(\lim_{s\to\infty} s \cdot I(s) = \lim_{s\to\infty} \frac{s \cdot V_s/s}{R+sL} = \lim_{s\to\infty} \frac{V_s}{R+sL} = 0\). So \(i(0^+)=0\). The inductor initially opposes any change in current.</li>
    <li>Final Value Theorem: \(\lim_{s\to 0} s \cdot I(s) = \lim_{s\to 0} \frac{V_s}{R+sL} = V_s/R = {@Vs/R@}\;\text{A}\). At steady state, the inductor acts as a short circuit.</li>
    <li>The time constant \(\tau = {@tau@}\;\text{s}\) is the time it takes for the current to reach approximately \(63.2\%\) of its final value. After \(5\tau\), the current is within \(1\%\) of its final value.</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* Variant 1: RL Step Response */
Vs: rand([10, 12, 15, 20]);
R: rand([100, 200, 500, 1000]);
L: rand([0.01, 0.02, 0.05, 0.1]);

/* Computed quantities */
tau: L/R;
Isteady: Vs/R;

/* Transfer function H(s) = I(s)/V(s) = 1/(R + s*L) */
Hs: 1/(R + s*L);

/* Part A correct answer */
correctA: a;

/* Part B correct answer */
correctB: 1/(R + s*L);

/* Part C correct answer */
correctC: tau;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V1 RL Step: Vs={@Vs@}, R={@R@}, L={@L@}, tau={@tau@}, H(s)={@Hs@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RL step response. Series RL driven by Vs*u(t). 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>algebraic</type>
        <tans>correctB</tans>
        <boxsize>25</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>1</insertstars>
        <syntaxhint>1/(R+s*L)</syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! An inductor with zero initial conditions is modeled as impedance \(sL\) in the s-domain.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. With zero initial conditions, the inductor \(L\) is modeled in the s-domain by its impedance \(Z_L = sL\). The correct answer is <strong>(a)</strong>.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check transfer function H(s)</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The transfer function is \(H(s) = \dfrac{1}{R + sL} = {@correctB@}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. By KVL in the s-domain: \(V(s) = I(s)(R + sL)\), so \(H(s) = \dfrac{I(s)}{V(s)} = \dfrac{1}{R+sL} = {@correctB@}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check time constant tau</description>
            <answertest>NumAbsolute</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.0001</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. For an RL circuit, the time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <deployedseed>12345</deployedseed>
    <deployedseed>23456</deployedseed>
    <deployedseed>34567</deployedseed>
    <deployedseed>45678</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 2: RC Step Response                                   -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V2_RC_Step_Response</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Q3 — Laplace Transform: RC Step-Response Circuit (13 pts)</h3>

<p>A series RC circuit consists of a resistor \(R = {@R@}\;\Omega\) and a capacitor \(C = {@C@}\;\text{F}\) connected in series. At \(t = 0\), a DC voltage source \(V_s = {@Vs@}\;\text{V}\) is applied as a step input \(V_s \cdot u(t)\). Assume <strong>zero initial conditions</strong> (capacitor initially uncharged).</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v2_rc_step.svg]</strong></p>

<hr/>

<h4>Part A (2 pts) — S-Domain Component Model (MCQ)</h4>
<p>When transforming the circuit to the s-domain with zero initial conditions, how is the capacitor \(C\) modeled?</p>
<p>Select the correct s-domain impedance of the capacitor:</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <p><strong>Options:</strong></p>
    <ul>
        <li>(a) \(sC\)</li>
        <li>(b) \(\dfrac{1}{sC}\)</li>
        <li>(c) \(\dfrac{C}{s}\)</li>
        <li>(d) \(\dfrac{s}{C}\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts) — Transfer Function</h4>
<p>Determine the transfer function relating the voltage across the capacitor to the input voltage:</p>
\[H(s) = \frac{V_C(s)}{V_s(s)}\]
<p>Express your answer as a rational function of \(s\). Use the variable <code>s</code>.</p>
<p>For example, if your answer were \(\dfrac{1}{1+0.5 s}\), type <code>1/(1+0.5*s)</code>.</p>
<p>\(H(s) = \) [[input:ansB]] [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts) — Time Constant</h4>
<p>Compute the time constant \(\tau\) of this RC circuit.</p>
<p>Give your answer in seconds as a decimal number (e.g., <code>0.047</code>).</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Physical Meaning of Time Constant (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>Use the <strong>Initial Value Theorem</strong> to find \(v_C(0^+)\).</li>
    <li>Use the <strong>Final Value Theorem</strong> to find \(v_C(\infty)\).</li>
    <li>Explain the <strong>physical meaning</strong> of the time constant \(\tau = {@tau@}\;\text{s}\) in this circuit. At \(t = \tau\), what percentage of the final voltage has the capacitor reached?</li>
</ul>
<p><strong>Write your answer below:</strong></p>
<p><textarea name="essayD_v2" rows="12" cols="80" style="width:100%; font-family:monospace;">
</textarea></p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution — Variant 2 (RC Step Response)</h4>
<p><strong>Part A:</strong> With zero initial conditions, a capacitor \(C\) in the s-domain is represented by impedance \(Z_C = \dfrac{1}{sC}\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> By voltage divider:
\[H(s) = \frac{V_C(s)}{V_s(s)} = \frac{1/(sC)}{R + 1/(sC)} = \frac{1}{1 + sRC} = {@correctB@}\]</p>
<p><strong>Part C:</strong> The time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>Initial Value Theorem: \(v_C(0^+) = \lim_{s\to\infty} s \cdot V_C(s) = \lim_{s\to\infty} \frac{s \cdot V_s/s}{1+sRC} = \lim_{s\to\infty} \frac{V_s}{1+sRC} = 0\). The capacitor voltage starts at zero.</li>
    <li>Final Value Theorem: \(v_C(\infty) = \lim_{s\to 0} s \cdot V_C(s) = \lim_{s\to 0} \frac{V_s}{1+sRC} = V_s = {@Vs@}\;\text{V}\). The capacitor charges fully to the source voltage.</li>
    <li>The time constant \(\tau = {@tau@}\;\text{s}\) represents the time at which the capacitor reaches approximately \(63.2\%\) of its final voltage (\(1 - e^{-1} \approx 0.632\)). After \(5\tau\), the capacitor is within \(1\%\) of its final value.</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* Variant 2: RC Step Response */
Vs: rand([5, 10, 12, 15]);
R: rand([1000, 2200, 3300, 4700]);
C: rand([10*10^(-6), 22*10^(-6), 47*10^(-6), 100*10^(-6)]);

/* Computed quantities */
tau: R*C;

/* Transfer function H(s) = Vc(s)/Vs(s) = 1/(1 + s*R*C) */
correctB: 1/(1 + s*R*C);

/* Part A correct answer */
correctA: b;

/* Part C correct answer */
correctC: tau;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V2 RC Step: Vs={@Vs@}, R={@R@}, C={@C@}, tau={@tau@}, H(s)={@correctB@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RC step response. Series RC driven by Vs*u(t). 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>algebraic</type>
        <tans>correctB</tans>
        <boxsize>25</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>1</insertstars>
        <syntaxhint>1/(1+s*R*C)</syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! A capacitor with zero initial conditions is modeled as impedance \(\dfrac{1}{sC}\) in the s-domain.</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. With zero initial conditions, the capacitor \(C\) is modeled in the s-domain by its impedance \(Z_C = \dfrac{1}{sC}\). The correct answer is <strong>(b)</strong>.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check transfer function H(s)</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The transfer function is \(H(s) = \dfrac{1}{1+sRC} = {@correctB@}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. Using voltage divider in the s-domain: \(H(s) = \dfrac{1/(sC)}{R+1/(sC)} = \dfrac{1}{1+sRC} = {@correctB@}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check time constant tau</description>
            <answertest>NumAbsolute</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.0001</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. For an RC circuit, the time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <deployedseed>56789</deployedseed>
    <deployedseed>67890</deployedseed>
    <deployedseed>78901</deployedseed>
    <deployedseed>89012</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 3: RL Natural Response                                -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V3_RL_Natural_Response</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Q3 — Laplace Transform: RL Natural Response Circuit (13 pts)</h3>

<p>An inductor \(L = {@L@}\;\text{H}\) carries an initial current \(I_0 = {@I0@}\;\text{A}\) at \(t = 0\). At \(t = 0\), the source is removed and the inductor is connected across a resistor \(R = {@R@}\;\Omega\). There is <strong>no external source</strong> — this is a natural (source-free) response.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v3_rl_natural.svg]</strong></p>

<hr/>

<h4>Part A (2 pts) — S-Domain Component Model (MCQ)</h4>
<p>In the s-domain, how is an inductor \(L\) with a nonzero initial current \(I_0\) modeled?</p>
<p>Select the correct s-domain model:</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <p><strong>Options:</strong></p>
    <ul>
        <li>(a) Impedance \(sL\) only (no initial condition source)</li>
        <li>(b) Impedance \(sL\) in series with a voltage source \(LI_0\) (representing the initial condition)</li>
        <li>(c) Impedance \(\dfrac{1}{sL}\) in parallel with a current source \(I_0/s\)</li>
        <li>(d) Impedance \(sL\) in series with a voltage source \(I_0/s\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts) — S-Domain Current Expression</h4>
<p>Using the s-domain circuit model with the initial condition, find the expression for \(I(s)\), the Laplace transform of the inductor current \(i(t)\).</p>
<p>Express your answer as a function of \(s\). Use the variable <code>s</code>.</p>
<p>For example, if your answer were \(\dfrac{2}{s+100}\), type <code>2/(s+100)</code>.</p>
<p>\(I(s) = \) [[input:ansB]] [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts) — Time Constant</h4>
<p>Compute the time constant \(\tau\) of this RL natural response.</p>
<p>Give your answer in seconds as a decimal number (e.g., <code>0.0002</code>).</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Energy Dissipation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>Use the <strong>Initial Value Theorem</strong> to verify that \(i(0^+) = I_0 = {@I0@}\;\text{A}\).</li>
    <li>Use the <strong>Final Value Theorem</strong> to verify that \(i(\infty) = 0\).</li>
    <li>The initial energy stored in the inductor is \(W_0 = \frac{1}{2}LI_0^2 = {@W0@}\;\text{J}\). Where does this energy go as \(t \to \infty\)? Explain the physical process.</li>
</ul>
<p><strong>Write your answer below:</strong></p>
<p><textarea name="essayD_v3" rows="12" cols="80" style="width:100%; font-family:monospace;">
</textarea></p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution — Variant 3 (RL Natural Response)</h4>
<p><strong>Part A:</strong> An inductor with initial current \(I_0\) is modeled in the s-domain as impedance \(sL\) in series with a voltage source \(LI_0\) (with polarity such that it drives current in the direction of \(I_0\)). This arises from \(\mathcal{L}\{L\,di/dt\} = sL \cdot I(s) - LI_0\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> KVL around the loop: \(sL \cdot I(s) - LI_0 + R \cdot I(s) = 0\), so:
\[I(s) = \frac{LI_0}{sL + R} = \frac{I_0}{s + R/L} = {@correctB@}\]</p>
<p><strong>Part C:</strong> The time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>Initial Value Theorem: \(\lim_{s\to\infty} s \cdot I(s) = \lim_{s\to\infty} \frac{s \cdot I_0}{s + R/L} = I_0 = {@I0@}\;\text{A}\). Confirmed.</li>
    <li>Final Value Theorem: \(\lim_{s\to 0} s \cdot I(s) = \lim_{s\to 0} \frac{s \cdot I_0}{s + R/L} = 0\). The current decays to zero.</li>
    <li>The initial energy \(W_0 = \frac{1}{2}({@L@})({@I0@})^2 = {@W0@}\;\text{J}\) is entirely dissipated as heat in the resistor \(R\). As the current decays exponentially, the instantaneous power \(p(t) = i^2(t)R\) is dissipated in the resistor, and integrating over all time gives \(\int_0^\infty i^2 R\,dt = W_0\).</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* Variant 3: RL Natural Response */
I0: rand([0.5, 1, 2, 3, 5]);
R: rand([50, 100, 200, 500]);
L: rand([0.01, 0.02, 0.05, 0.1]);

/* Computed quantities */
tau: L/R;
W0: (1/2)*L*I0^2;

/* I(s) = I0/(s + R/L) = L*I0/(s*L + R) */
correctB: I0/(s + R/L);

/* Part A correct answer */
correctA: b;

/* Part C correct answer */
correctC: tau;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V3 RL Natural: I0={@I0@}, R={@R@}, L={@L@}, tau={@tau@}, I(s)={@correctB@}, W0={@W0@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RL natural response. Inductor with initial current, resistor. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>algebraic</type>
        <tans>correctB</tans>
        <boxsize>25</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>1</insertstars>
        <syntaxhint>I0/(s+R/L)</syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! An inductor with initial current \(I_0\) is modeled as impedance \(sL\) in series with a voltage source \(LI_0\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. The s-domain model of an inductor with initial current \(I_0\) includes impedance \(sL\) in series with a voltage source \(LI_0\). This comes from \(\mathcal{L}\{L\,di/dt\} = sL \cdot I(s) - LI_0\). The correct answer is <strong>(b)</strong>.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check I(s) expression</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! \(I(s) = \dfrac{I_0}{s + R/L} = {@correctB@}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. From KVL: \(sL \cdot I(s) - LI_0 + R \cdot I(s) = 0\), giving \(I(s) = \dfrac{LI_0}{sL+R} = \dfrac{I_0}{s+R/L} = {@correctB@}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check time constant tau</description>
            <answertest>NumAbsolute</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.0001</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. For an RL circuit, the time constant is \(\tau = L/R = {@L@}/{@R@} = {@tau@}\;\text{s}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <deployedseed>11111</deployedseed>
    <deployedseed>22222</deployedseed>
    <deployedseed>33333</deployedseed>
    <deployedseed>44444</deployedseed>
</question>

<!-- ============================================================ -->
<!-- VARIANT 4: RC Natural Response                                -->
<!-- ============================================================ -->
<question type="stack">
    <name>
        <text>Q3_Medium_B_V4_RC_Natural_Response</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[
<h3>Q3 — Laplace Transform: RC Natural Response Circuit (13 pts)</h3>

<p>A capacitor \(C = {@C@}\;\text{F}\) is initially charged to a voltage \(V_0 = {@V0@}\;\text{V}\). At \(t = 0\), the capacitor is connected across a resistor \(R = {@R@}\;\Omega\). There is <strong>no external source</strong> — this is a natural (source-free) response.</p>

<p><strong>[INSERT DIAGRAM: diagrams/q3/q3_v4_rc_natural.svg]</strong></p>

<hr/>

<h4>Part A (2 pts) — S-Domain Component Model (MCQ)</h4>
<p>In the s-domain, how is a capacitor \(C\) with an initial voltage \(V_0\) modeled?</p>
<p>Select the correct s-domain model:</p>
<p>[[input:ansA]] [[validation:ansA]]</p>
<div style="margin-left:20px;">
    <p><strong>Options:</strong></p>
    <ul>
        <li>(a) Impedance \(\dfrac{1}{sC}\) only (no initial condition source)</li>
        <li>(b) Impedance \(\dfrac{1}{sC}\) in series with a voltage source \(\dfrac{V_0}{s}\) (representing the initial condition)</li>
        <li>(c) Impedance \(sC\) in series with a voltage source \(V_0 \cdot s\)</li>
        <li>(d) Impedance \(\dfrac{1}{sC}\) in parallel with a current source \(\dfrac{V_0}{R}\)</li>
    </ul>
</div>
<p>Enter the letter of your choice: <strong>a</strong>, <strong>b</strong>, <strong>c</strong>, or <strong>d</strong>.</p>
[[feedback:prtA]]

<hr/>

<h4>Part B (4 pts) — S-Domain Voltage Expression</h4>
<p>Using the s-domain circuit model with the initial condition, find the expression for \(V_C(s)\), the Laplace transform of the capacitor voltage \(v_C(t)\).</p>
<p>Express your answer as a function of \(s\). Use the variable <code>s</code>.</p>
<p>For example, if your answer were \(\dfrac{10}{s+50}\), type <code>10/(s+50)</code>.</p>
<p>\(V_C(s) = \) [[input:ansB]] [[validation:ansB]]</p>
[[feedback:prtB]]

<hr/>

<h4>Part C (4 pts) — Time Constant</h4>
<p>Compute the time constant \(\tau\) of this RC natural response.</p>
<p>Give your answer in seconds as a decimal number (e.g., <code>0.047</code>).</p>
<p>\(\tau = \) [[input:ansC]] seconds [[validation:ansC]]</p>
[[feedback:prtC]]

<hr/>

<h4>Part D (3 pts) — Initial/Final Value Check &amp; Energy Dissipation (Essay)</h4>
<p>Answer the following:</p>
<ul>
    <li>Use the <strong>Initial Value Theorem</strong> to verify that \(v_C(0^+) = V_0 = {@V0@}\;\text{V}\).</li>
    <li>Use the <strong>Final Value Theorem</strong> to verify that \(v_C(\infty) = 0\).</li>
    <li>The initial energy stored in the capacitor is \(W_0 = \frac{1}{2}CV_0^2 = {@W0@}\;\text{J}\). Explain where this energy goes as \(t \to \infty\) and describe the physical process of energy dissipation.</li>
</ul>
<p><strong>Write your answer below:</strong></p>
<p><textarea name="essayD_v4" rows="12" cols="80" style="width:100%; font-family:monospace;">
</textarea></p>
]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text><![CDATA[
<h4>Solution — Variant 4 (RC Natural Response)</h4>
<p><strong>Part A:</strong> A capacitor with initial voltage \(V_0\) is modeled in the s-domain as impedance \(\dfrac{1}{sC}\) in series with a voltage source \(\dfrac{V_0}{s}\). This comes from the Laplace transform of the capacitor constitutive relation: \(V_C(s) = \dfrac{1}{sC}I(s) + \dfrac{V_0}{s}\). Answer: <strong>(b)</strong>.</p>
<p><strong>Part B:</strong> The current flowing through the resistor is \(I(s)\), and KVL gives:
\[\frac{V_0}{s} + \frac{1}{sC}I(s) = R \cdot I(s) \quad \Rightarrow \quad \text{(or using node voltage directly)}\]
Actually, the capacitor voltage satisfies \(V_C(s) = \dfrac{V_0}{s + 1/(RC)}\) because the discharge equation transforms to:
\[RC \cdot sV_C(s) - RC \cdot V_0 + V_C(s) = 0 \implies V_C(s)(sRC+1) = RCV_0 \implies V_C(s) = \frac{V_0}{s+1/(RC)} = {@correctB@}\]</p>
<p><strong>Part C:</strong> The time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>
<p><strong>Part D:</strong></p>
<ul>
    <li>Initial Value Theorem: \(\lim_{s\to\infty} s \cdot V_C(s) = \lim_{s\to\infty} \frac{s \cdot V_0}{s + 1/(RC)} = V_0 = {@V0@}\;\text{V}\). Confirmed.</li>
    <li>Final Value Theorem: \(\lim_{s\to 0} s \cdot V_C(s) = \lim_{s\to 0} \frac{s \cdot V_0}{s + 1/(RC)} = 0\). The voltage decays to zero.</li>
    <li>The initial energy \(W_0 = \frac{1}{2}({@C@})({@V0@})^2 = {@W0@}\;\text{J}\) is entirely dissipated as heat in the resistor \(R\). As the capacitor discharges, current flows through \(R\), and the power \(p(t) = v_C^2(t)/R\) is converted to thermal energy. The total energy dissipated is \(\int_0^\infty v_C^2/R\,dt = W_0\).</li>
</ul>
]]></text>
    </generalfeedback>
    <defaultgrade>13</defaultgrade>
    <penalty>0.1</penalty>
    <hidden>0</hidden>
    <idnumber></idnumber>
    <stackversion>
        <text>2024120900</text>
    </stackversion>
    <questionvariables>
        <text><![CDATA[/* Variant 4: RC Natural Response */
V0: rand([5, 10, 12, 15, 20]);
R: rand([1000, 2200, 4700, 10000]);
C: rand([10*10^(-6), 22*10^(-6), 47*10^(-6), 100*10^(-6)]);

/* Computed quantities */
tau: R*C;
W0: (1/2)*C*V0^2;

/* Vc(s) = V0/(s + 1/(R*C)) */
correctB: V0/(s + 1/(R*C));

/* Part A correct answer */
correctA: b;

/* Part C correct answer */
correctC: tau;
]]></text>
    </questionvariables>
    <specificfeedback format="html">
        <text></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[V4 RC Natural: V0={@V0@}, R={@R@}, C={@C@}, tau={@tau@}, Vc(s)={@correctB@}, W0={@W0@}]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[Laplace Transform — RC natural response. Capacitor with initial voltage, resistor. 13 pts total.]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text><![CDATA[<span style="color:green;"><strong>Correct!</strong></span>]]></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text><![CDATA[<span style="color:orange;"><strong>Partially correct.</strong></span>]]></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text><![CDATA[<span style="color:red;"><strong>Incorrect.</strong></span>]]></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ansA</name>
        <type>algebraic</type>
        <tans>correctA</tans>
        <boxsize>5</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords>a,b,c,d</allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>3</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansB</name>
        <type>algebraic</type>
        <tans>correctB</tans>
        <boxsize>25</boxsize>
        <strictsyntax>0</strictsyntax>
        <insertstars>1</insertstars>
        <syntaxhint>V0/(s+1/(R*C))</syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <input>
        <name>ansC</name>
        <type>numerical</type>
        <tans>correctC</tans>
        <boxsize>15</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>0</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>1</showvalidation>
        <options></options>
    </input>
    <prt>
        <name>prtA</name>
        <value>2</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check MCQ answer for Part A</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansA</sans>
            <tans>correctA</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtA-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! A capacitor with initial voltage \(V_0\) is modeled as impedance \(\dfrac{1}{sC}\) in series with a voltage source \(\dfrac{V_0}{s}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtA-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. The s-domain model of a capacitor with initial voltage \(V_0\) includes impedance \(\dfrac{1}{sC}\) in series with a voltage source \(\dfrac{V_0}{s}\). This arises from \(V_C(s) = \dfrac{1}{sC}I(s) + \dfrac{V_0}{s}\). The correct answer is <strong>(b)</strong>.</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtB</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check Vc(s) expression</description>
            <answertest>AlgEquiv</answertest>
            <sans>ansB</sans>
            <tans>correctB</tans>
            <testoptions></testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtB-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! \(V_C(s) = \dfrac{V_0}{s + 1/(RC)} = {@correctB@}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtB-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. The RC discharge transforms to \(V_C(s)(sRC+1) = RCV_0\), giving \(V_C(s) = \dfrac{V_0}{s+1/(RC)} = {@correctB@}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <prt>
        <name>prtC</name>
        <value>4</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description>Check time constant tau</description>
            <answertest>NumAbsolute</answertest>
            <sans>ansC</sans>
            <tans>correctC</tans>
            <testoptions>0.0001</testoptions>
            <quiet>0</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prtC-1-T</trueanswernote>
            <truefeedback format="html">
                <text><![CDATA[<p>Correct! The time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prtC-1-F</falseanswernote>
            <falsefeedback format="html">
                <text><![CDATA[<p>Incorrect. For an RC circuit, the time constant is \(\tau = RC = {@R@} \times {@C@} = {@tau@}\;\text{s}\).</p>]]></text>
            </falsefeedback>
        </node>
    </prt>
    <deployedseed>55555</deployedseed>
    <deployedseed>66666</deployedseed>
    <deployedseed>77777</deployedseed>
    <deployedseed>88888</deployedseed>
</question>

</quiz>
