<?xml version="1.0" encoding="UTF-8"?>
<quiz>

<!-- ============================================================ -->
<!-- Week 10 - Q1: Series RLC Natural Response (Switch Opening)    -->
<!-- Course: Electromagnetism and Circuit Analysis, LUT University  -->
<!-- ============================================================ -->
<question type="stack">
  <name>
    <text>Week 10 – Q1: Series RLC Natural Response (Switch Opening)</text>
  </name>
  <questiontext format="html">
    <text><![CDATA[
<h3>Week 10 &ndash; Q1: Series RLC Natural Response (Switch Opening)</h3>
<p><em>Course: Electromagnetism and Circuit Analysis &mdash; LUT University, Finland</em></p>
<p><em>This is a training exercise &mdash; you may use books, tutors, and AI. The question requires multi-step reasoning.</em></p>
<p><strong>Total: 8 points | Penalty per wrong attempt: 10%</strong></p>

<hr />

<p>A series RLC circuit is powered by a DC voltage source \(V_s = {@Vs_val@}\;\text{V}\).
The circuit contains a resistor \(R = {@R_val@}\;\Omega\), an inductor \(L = {@L_val@}\;\text{H}\),
and a capacitor \(C = {@C_uF@}\;\mu\text{F}\), all connected in series with the source and a switch.</p>

<p>The circuit has been in <strong>steady state for a long time</strong>.
At \(t = 0\), the switch <strong>opens</strong>, disconnecting the DC source and leaving the series RLC loop isolated.</p>

<p>Your task: determine the complete natural response \(i(t)\) for \(t &gt; 0\) by working through parts (a)&ndash;(f) below.</p>

<hr />

<h4>Part (a) &mdash; DC Steady-State Concept (1 pt)</h4>
<p>Before the switch opens (\(t &lt; 0\)), the circuit has been in DC steady state for a very long time.
Which statement correctly describes the initial conditions?</p>
<ol>
  <li>Capacitor = open circuit at DC &rarr; \(i(0^-) = 0\;\text{A}\), \(v_C(0^-) = V_s\)</li>
  <li>Inductor = open circuit at DC &rarr; \(i(0^-) = 0\), \(v_C(0^-) = 0\)</li>
  <li>Both components block DC &rarr; \(i = 0\), \(v_C = 0\)</li>
  <li>Inductor = short circuit carrying \(V_s/R\) &rarr; \(v_C = 0\)</li>
</ol>
<p>Enter the number of the correct option (1&ndash;4):
[[input:ans1]] [[validation:ans1]]</p>

<hr />

<h4>Part (b1) &mdash; Initial Current \(i(0^+)\) (1 pt)</h4>
<p>Using the DC steady-state analysis and continuity conditions, what is \(i(0^+)\) in Amperes?</p>
<p>\(i(0^+) = \) [[input:ans2]] A [[validation:ans2]]</p>

<hr />

<h4>Part (b2) &mdash; Initial Capacitor Voltage \(v_C(0^+)\) (1 pt)</h4>
<p>What is the value of \(v_C(0^+)\) in Volts? (Express to 3 significant figures.)</p>
<p>\(v_C(0^+) = \) [[input:ans3]] V [[validation:ans3]]</p>

<hr />

<h4>Part (c1) &mdash; Damping Coefficient \(\alpha\) (1 pt)</h4>
<p>Calculate the damping coefficient \(\alpha\) (in rad/s) for this series RLC circuit.
(Express to 3 significant figures.)</p>
<p>\(\alpha = \) [[input:ans4]] rad/s [[validation:ans4]]</p>

<hr />

<h4>Part (c2) &mdash; Natural Frequency \(\omega_0\) (1 pt)</h4>
<p>Calculate the resonant (natural) frequency \(\omega_0\) (in rad/s).
(Express to 3 significant figures.)</p>
<p>\(\omega_0 = \) [[input:ans5]] rad/s [[validation:ans5]]</p>

<hr />

<h4>Part (d) &mdash; Damping Regime Classification (1 pt)</h4>
<p>Based on your values of \(\alpha\) and \(\omega_0\), classify the damping regime:</p>
<ol>
  <li>Overdamped (\(\alpha &gt; \omega_0\))</li>
  <li>Underdamped (\(\alpha &lt; \omega_0\))</li>
  <li>Critically damped (\(\alpha = \omega_0\))</li>
</ol>
<p>Enter the number of the correct regime (1&ndash;3):
[[input:ans6]] [[validation:ans6]]</p>

<hr />

<h4>Part (e) &mdash; Characteristic Root \(s_1\) (1 pt)</h4>
<p>Find the characteristic root \(s_1\) with the <strong>larger real part</strong>
(or positive imaginary part if complex conjugate pair).</p>
<p><em>Syntax hint: For complex roots use <code>j</code> for the imaginary unit,
e.g., <code>-2800+9600*j</code></em></p>
<p>\(s_1 = \) [[input:ans7]] [[validation:ans7]]</p>

<hr />

<h4>Part (f) &mdash; Complete Current Response \(i(t)\) (1 pt)</h4>
<p>Write the complete expression for \(i(t)\) (in Amperes) for \(t &gt; 0\).</p>
<p><em>Syntax hint: Use <code>exp(...)</code>, <code>sin(...)</code>, <code>cos(...)</code>,
and <code>t</code> as the time variable.<br />
Example: <code>-0.104*exp(-2800*t)*sin(9600*t)</code></em></p>
<p>\(i(t) = \) [[input:ans8]] A [[validation:ans8]]</p>
]]></text>
  </questiontext>

  <generalfeedback format="html">
    <text><![CDATA[
<h3>Worked Solution</h3>

<h4>Step 1: DC Steady-State Analysis</h4>
<p>After a very long time in DC steady state:</p>
<ul>
  <li>The <strong>capacitor</strong> acts as an <strong>open circuit</strong> (it blocks DC current after fully charging).</li>
  <li>The <strong>inductor</strong> acts as a <strong>short circuit</strong> (zero voltage drop across it at constant current).</li>
</ul>
<p>Since the capacitor is open, no current flows in the series loop:</p>
<p>\[i(0^-) = 0\;\text{A}\]</p>
<p>With no current flowing, there is no voltage drop across \(R\) or \(L\).
All the source voltage appears across the capacitor:</p>
<p>\[v_C(0^-) = V_s = {@Vs_val@}\;\text{V}\]</p>
<p>By continuity of inductor current and capacitor voltage:</p>
<p>\[i(0^+) = i(0^-) = 0\;\text{A}, \quad v_C(0^+) = v_C(0^-) = {@ta3@}\;\text{V}\]</p>

<h4>Step 2: Damping Parameters</h4>
<p>For a <strong>series</strong> RLC circuit:</p>
<p>\[\alpha = \frac{R}{2L} = \frac{{@R_val@}}{2 \times {@L_val@}} = {@ta_alpha@}\;\text{rad/s}\]</p>
<p>\[\omega_0 = \frac{1}{\sqrt{LC}} = \frac{1}{\sqrt{{@L_val@} \times {@C_val@}}} = {@ta_omega0@}\;\text{rad/s}\]</p>

<h4>Step 3: Damping Regime Classification</h4>
<p>Comparing \(\alpha = {@ta_alpha@}\;\text{rad/s}\) with \(\omega_0 = {@ta_omega0@}\;\text{rad/s}\):</p>
<p>The circuit is <strong>{@regime_name@}</strong>.</p>

<h4>Step 4: Initial Derivative</h4>
<p>After the switch opens, KVL around the isolated series RLC loop gives:</p>
<p>\[L\frac{di}{dt} + Ri + v_C = 0\]</p>
<p>At \(t = 0^+\):</p>
<p>\[\frac{di}{dt}(0^+) = \frac{1}{L}\left(-R \cdot i(0^+) - v_C(0^+)\right)
= \frac{1}{{@L_val@}}\left(-{@R_val@} \cdot 0 - {@Vs_val@}\right)
= {@didt0@}\;\text{A/s}\]</p>

<h4>Step 5: Complete Response</h4>
<p>Applying the initial conditions \(i(0) = 0\) and \(\frac{di}{dt}(0) = {@didt0@}\;\text{A/s}\):</p>
<p>\[i(t) = {@ta8@}\;\text{A}, \quad t &gt; 0\]</p>

<hr />

<h4>Answer Summary</h4>
<table border="1" cellpadding="5" cellspacing="0">
  <tr><th>Part</th><th>Quantity</th><th>Correct Answer</th></tr>
  <tr><td>(a)</td><td>DC steady-state concept</td><td>Option {@ta1@}</td></tr>
  <tr><td>(b1)</td><td>\(i(0^+)\)</td><td>{@ta2@} A</td></tr>
  <tr><td>(b2)</td><td>\(v_C(0^+)\)</td><td>{@ta3@} V</td></tr>
  <tr><td>(c1)</td><td>\(\alpha\)</td><td>{@ta4@} rad/s</td></tr>
  <tr><td>(c2)</td><td>\(\omega_0\)</td><td>{@ta5@} rad/s</td></tr>
  <tr><td>(d)</td><td>Damping regime</td><td>{@regime_name@} (option {@ta6@})</td></tr>
  <tr><td>(e)</td><td>\(s_1\)</td><td>{@ta7@}</td></tr>
  <tr><td>(f)</td><td>\(i(t)\)</td><td>{@ta8@} A</td></tr>
</table>
]]></text>
  </generalfeedback>

  <defaultgrade>8.0000000</defaultgrade>
  <penalty>0.1000000</penalty>
  <hidden>0</hidden>
  <idnumber/>

  <stackversion>
    <text>2024032101</text>
  </stackversion>

  <questionvariables>
    <text><![CDATA[
/* ===== Parameter set selection ===== */
set: rand(3)+1;

/* ===== Component values ===== */
R_val:  if set=1 then 560  else if set=2 then 3000 else 2000;
L_val:  0.1;
C_val:  if set=1 then 1e-7 else if set=2 then 1e-6 else 1e-7;
C_uF:   float(C_val*1e6);
Vs_val: if set=1 then 100  else if set=2 then 50   else 75;

/* ===== Damping parameters ===== */
ta_alpha:  float(R_val/(2*L_val));
ta_omega0: float(1/sqrt(L_val*C_val));
disc: float(ta_alpha^2 - ta_omega0^2);

/* ===== Regime classification ===== */
/* 1=overdamped, 2=underdamped, 3=critically damped */
ta_regime: if abs(ta_alpha-ta_omega0) < 2.0 then 3
           else if disc > 0 then 1
           else 2;
regime_name: if ta_regime=1 then "overdamped"
             else if ta_regime=2 then "underdamped"
             else "critically damped";

/* ===== Damped frequency / root separation ===== */
ta_omegad: float(sqrt(abs(disc)));

/* ===== Characteristic roots ===== */
s1_r: if ta_regime=1 then float(-ta_alpha+sqrt(disc)) else float(-ta_alpha);
s2_r: if ta_regime=1 then float(-ta_alpha-sqrt(disc)) else float(-ta_alpha);

/* ===== Initial derivative ===== */
didt0: float(-Vs_val/L_val);

/* ===== Coefficients for i(t) ===== */
ta_A1: if ta_regime=1 then float(didt0/(s1_r-s2_r)) else 0;
ta_A2: if ta_regime=1 then float(-ta_A1) else 0;
ta_B2: if ta_regime=2 then float(didt0/ta_omegad) else 0;
ta_D1: if ta_regime=3 then float(didt0) else 0;

/* ===== Teacher answers ===== */
ta1: 1;
ta2: 0;
ta3: float(Vs_val);
ta4: ta_alpha;
ta5: ta_omega0;
ta6: ta_regime;

ta7: if ta_regime=1 then s1_r
     else if ta_regime=2 then (-ta_alpha + ta_omegad*%i)
     else float(-ta_alpha);

ta8: if ta_regime=1 then ta_A1*exp(s1_r*t)+ta_A2*exp(s2_r*t)
     else if ta_regime=2 then ta_B2*exp(-ta_alpha*t)*sin(ta_omegad*t)
     else ta_D1*t*exp(-ta_alpha*t);
]]></text>
  </questionvariables>

  <specificfeedback format="html">
    <text><![CDATA[
<h4>Part (a) &mdash; DC Steady-State Concept</h4>
<p>Your answer: {@ans1@} &nbsp;|&nbsp; Correct answer: {@ta1@}</p>
<div>[[feedback:prt1]]</div>

<h4>Part (b1) &mdash; Initial Current \(i(0^+)\)</h4>
<p>Your answer: {@ans2@} &nbsp;|&nbsp; Correct answer: {@ta2@}</p>
<div>[[feedback:prt2]]</div>

<h4>Part (b2) &mdash; Initial Capacitor Voltage \(v_C(0^+)\)</h4>
<p>Your answer: {@ans3@} &nbsp;|&nbsp; Correct answer: {@ta3@} V</p>
<div>[[feedback:prt3]]</div>

<h4>Part (c1) &mdash; Damping Coefficient \(\alpha\)</h4>
<p>Your answer: {@ans4@} &nbsp;|&nbsp; Correct answer: {@ta4@} rad/s</p>
<div>[[feedback:prt4]]</div>

<h4>Part (c2) &mdash; Natural Frequency \(\omega_0\)</h4>
<p>Your answer: {@ans5@} &nbsp;|&nbsp; Correct answer: {@ta5@} rad/s</p>
<div>[[feedback:prt5]]</div>

<h4>Part (d) &mdash; Damping Regime</h4>
<p>Your answer: {@ans6@} &nbsp;|&nbsp; Correct answer: {@ta6@} ({@regime_name@})</p>
<div>[[feedback:prt6]]</div>

<h4>Part (e) &mdash; Characteristic Root \(s_1\)</h4>
<p>Your answer: {@ans7@} &nbsp;|&nbsp; Correct answer: {@ta7@}</p>
<div>[[feedback:prt7]]</div>

<h4>Part (f) &mdash; Complete Response \(i(t)\)</h4>
<p>Your answer: {@ans8@} &nbsp;|&nbsp; Correct answer: {@ta8@}</p>
<div>[[feedback:prt8]]</div>
]]></text>
  </specificfeedback>

  <questionnote>
    <text>Set {@set@}: R={@R_val@} Ω, C={@C_uF@} µF, Vs={@Vs_val@} V | α={@ta_alpha@}, ω₀={@ta_omega0@} | {@regime_name@}</text>
  </questionnote>

  <questiondescription format="html">
    <text><![CDATA[<p>Series RLC Natural Response — Derive initial conditions from DC steady state, classify damping regime, and find i(t). Week 10 training exercise.</p>]]></text>
  </questiondescription>

  <questionsimplify>1</questionsimplify>
  <assumepositive>0</assumepositive>
  <assumereal>0</assumereal>

  <prtcorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:green;"><i class="fa fa-check"></i></span> Correct answer, well done!]]></text>
  </prtcorrect>
  <prtpartiallycorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:orange;"><i class="fa fa-adjust"></i></span> Your answer is partially correct.]]></text>
  </prtpartiallycorrect>
  <prtincorrect format="html">
    <text><![CDATA[<span style="font-size: 1.5em; color:red;"><i class="fa fa-times"></i></span> Incorrect answer.]]></text>
  </prtincorrect>

  <decimals>.</decimals>
  <scientificnotation>*10</scientificnotation>
  <multiplicationsign>dot</multiplicationsign>
  <sqrtsign>1</sqrtsign>
  <complexno>j</complexno>
  <inversetrig>cos-1</inversetrig>
  <logicsymbol>lang</logicsymbol>
  <matrixparens>[</matrixparens>
  <variantsselectionseed/>

  <!-- ============================================================ -->
  <!-- INPUTS (8 total)                                              -->
  <!-- ============================================================ -->

  <!-- ans1: Part (a) MCQ — DC steady-state concept -->
  <input>
    <name>ans1</name>
    <type>algebraic</type>
    <tans>ta1</tans>
    <boxsize>3</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>1</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans2: Part (b1) — i(0+) -->
  <input>
    <name>ans2</name>
    <type>numerical</type>
    <tans>ta2</tans>
    <boxsize>8</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans3: Part (b2) — vC(0+) -->
  <input>
    <name>ans3</name>
    <type>numerical</type>
    <tans>ta3</tans>
    <boxsize>10</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans4: Part (c1) — alpha -->
  <input>
    <name>ans4</name>
    <type>numerical</type>
    <tans>ta4</tans>
    <boxsize>12</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans5: Part (c2) — omega0 -->
  <input>
    <name>ans5</name>
    <type>numerical</type>
    <tans>ta5</tans>
    <boxsize>12</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans6: Part (d) MCQ — damping regime -->
  <input>
    <name>ans6</name>
    <type>algebraic</type>
    <tans>ta6</tans>
    <boxsize>3</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>0</insertstars>
    <syntaxhint/>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>1</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans7: Part (e) — characteristic root s1 -->
  <input>
    <name>ans7</name>
    <type>algebraic</type>
    <tans>ta7</tans>
    <boxsize>22</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>1</insertstars>
    <syntaxhint>-2800+9600*j</syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ans8: Part (f) — complete i(t) expression -->
  <input>
    <name>ans8</name>
    <type>algebraic</type>
    <tans>ta8</tans>
    <boxsize>45</boxsize>
    <strictsyntax>1</strictsyntax>
    <insertstars>1</insertstars>
    <syntaxhint>-0.104*exp(-2800*t)*sin(9600*t)</syntaxhint>
    <syntaxattribute>0</syntaxattribute>
    <forbidwords/>
    <allowwords/>
    <forbidfloat>0</forbidfloat>
    <requirelowestterms>0</requirelowestterms>
    <checkanswertype>0</checkanswertype>
    <mustverify>1</mustverify>
    <showvalidation>3</showvalidation>
    <options/>
  </input>

  <!-- ============================================================ -->
  <!-- PRTs (8 total, each worth 1 mark)                            -->
  <!-- ============================================================ -->

  <!-- prt1: Part (a) — DC steady-state MCQ -->
  <prt>
    <name>prt1</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check Part (a) MCQ answer</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans1</sans>
      <tans>ta1</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt1-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! In DC steady state, the capacitor is an open circuit (blocks DC current), so \(i(0^-) = 0\;\text{A}\) and all source voltage appears across the capacitor: \(v_C(0^-) = V_s = {@ta3@}\;\text{V}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt1-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The correct answer is {@ta1@}. In DC steady state: the <strong>capacitor</strong> acts as an <strong>open circuit</strong> (it fully charges and blocks DC). The <strong>inductor</strong> acts as a <strong>short circuit</strong> (zero voltage for constant current). Since the capacitor is open, no current flows: \(i(0^-) = 0\;\text{A}\), and \(v_C(0^-) = V_s = {@ta3@}\;\text{V}\).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt2: Part (b1) — i(0+) -->
  <prt>
    <name>prt2</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check i(0+) value</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans2</sans>
      <tans>ta2</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt2-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! By continuity of inductor current: \(i(0^+) = i(0^-) = {@ta2@}\;\text{A}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt2-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. Since no DC current flows through the series circuit at steady state (capacitor is open), \(i(0^-) = 0\;\text{A}\). By continuity of inductor current, \(i(0^+) = {@ta2@}\;\text{A}\).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt3: Part (b2) — vC(0+), 2-node cascade -->
  <prt>
    <name>prt3</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>NumRelative check vC(0+) within 5%</description>
      <answertest>NumRelative</answertest>
      <sans>ans3</sans>
      <tans>ta3</tans>
      <testoptions>0.05</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.7000000</truescore>
      <truepenalty/>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt3-0-T</trueanswernote>
      <truefeedback format="html">
        <text/>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt3-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The capacitor voltage at DC steady state equals the source voltage. By continuity: \(v_C(0^+) = V_s = {@ta3@}\;\text{V}\).</p>]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>SigFigsStrict check — 3 significant figures</description>
      <answertest>SigFigsStrict</answertest>
      <sans>ans3</sans>
      <tans>ta3</tans>
      <testoptions>3</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt3-1-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(v_C(0^+) = {@ta3@}\;\text{V}\), with proper significant figures.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.7000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt3-1-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Your value is numerically correct but not expressed to 3 significant figures. Score: 70%. Expected: {@ta3@} V.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt4: Part (c1) — alpha, 2-node cascade -->
  <prt>
    <name>prt4</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>NumRelative check alpha within 5%</description>
      <answertest>NumRelative</answertest>
      <sans>ans4</sans>
      <tans>ta4</tans>
      <testoptions>0.05</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.7000000</truescore>
      <truepenalty/>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt4-0-T</trueanswernote>
      <truefeedback format="html">
        <text/>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt4-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. For a <strong>series</strong> RLC circuit, the damping coefficient is \(\alpha = \frac{R}{2L}\), <strong>not</strong> \(\frac{1}{2RC}\) (which applies to <em>parallel</em> RLC). With \(R = {@R_val@}\;\Omega\) and \(L = {@L_val@}\;\text{H}\): \(\alpha = {@ta4@}\;\text{rad/s}\).</p>]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>SigFigsStrict check — 3 significant figures</description>
      <answertest>SigFigsStrict</answertest>
      <sans>ans4</sans>
      <tans>ta4</tans>
      <testoptions>3</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt4-1-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(\alpha = {@ta4@}\;\text{rad/s}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.7000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt4-1-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Your value is numerically correct but not expressed to 3 significant figures. Score: 70%. Expected: {@ta4@} rad/s.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt5: Part (c2) — omega0, 2-node cascade -->
  <prt>
    <name>prt5</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>NumRelative check omega0 within 5%</description>
      <answertest>NumRelative</answertest>
      <sans>ans5</sans>
      <tans>ta5</tans>
      <testoptions>0.05</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>0.7000000</truescore>
      <truepenalty/>
      <truenextnode>1</truenextnode>
      <trueanswernote>prt5-0-T</trueanswernote>
      <truefeedback format="html">
        <text/>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt5-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The natural frequency is \(\omega_0 = \frac{1}{\sqrt{LC}}\). With \(L = {@L_val@}\;\text{H}\) and \(C = {@C_val@}\;\text{F}\): \(\omega_0 = {@ta5@}\;\text{rad/s}\).</p>]]></text>
      </falsefeedback>
    </node>
    <node>
      <name>1</name>
      <description>SigFigsStrict check — 3 significant figures</description>
      <answertest>SigFigsStrict</answertest>
      <sans>ans5</sans>
      <tans>ta5</tans>
      <testoptions>3</testoptions>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt5-1-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(\omega_0 = {@ta5@}\;\text{rad/s}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.7000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt5-1-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Your value is numerically correct but not expressed to 3 significant figures. Score: 70%. Expected: {@ta5@} rad/s.</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt6: Part (d) — damping regime MCQ -->
  <prt>
    <name>prt6</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text><![CDATA[
alpha_sq: float(ta_alpha^2);
omega0_sq: float(ta_omega0^2);
]]></text>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check regime classification</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans6</sans>
      <tans>ta6</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt6-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! The circuit is <strong>{@regime_name@}</strong>.</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt6-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. Compare: \(\alpha^2 = {@alpha_sq@}\) vs \(\omega_0^2 = {@omega0_sq@}\). Since \(\alpha = {@ta_alpha@}\) and \(\omega_0 = {@ta_omega0@}\), the circuit is <strong>{@regime_name@}</strong> (option {@ta6@}).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt7: Part (e) — characteristic root s1 -->
  <prt>
    <name>prt7</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check s1 root</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans7</sans>
      <tans>ta7</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt7-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(s_1 = {@ta7@}\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt7-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. The characteristic root with the larger real part is \(s_1 = {@ta7@}\). Remember the characteristic equation is \(s^2 + \frac{R}{L}s + \frac{1}{LC} = 0\), giving \(s = -\alpha \pm \sqrt{\alpha^2 - \omega_0^2}\).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- prt8: Part (f) — complete i(t) expression -->
  <prt>
    <name>prt8</name>
    <value>1.0000000</value>
    <autosimplify>1</autosimplify>
    <feedbackstyle>1</feedbackstyle>
    <feedbackvariables>
      <text/>
    </feedbackvariables>
    <node>
      <name>0</name>
      <description>Check complete i(t) expression</description>
      <answertest>AlgEquiv</answertest>
      <sans>ans8</sans>
      <tans>ta8</tans>
      <testoptions/>
      <quiet>0</quiet>
      <truescoremode>=</truescoremode>
      <truescore>1.0000000</truescore>
      <truepenalty/>
      <truenextnode>-1</truenextnode>
      <trueanswernote>prt8-0-T</trueanswernote>
      <truefeedback format="html">
        <text><![CDATA[<p>Correct! \(i(t) = {@ta8@}\;\text{A}\) for \(t > 0\).</p>]]></text>
      </truefeedback>
      <falsescoremode>=</falsescoremode>
      <falsescore>0.0000000</falsescore>
      <falsepenalty/>
      <falsenextnode>-1</falsenextnode>
      <falseanswernote>prt8-0-F</falseanswernote>
      <falsefeedback format="html">
        <text><![CDATA[<p>Incorrect. Remember to apply <strong>both</strong> initial conditions: \(i(0) = 0\) and \(\frac{di}{dt}(0) = {@didt0@}\;\text{A/s}\). The complete response is \(i(t) = {@ta8@}\;\text{A}\).</p>]]></text>
      </falsefeedback>
    </node>
  </prt>

  <!-- ============================================================ -->
  <!-- DEPLOYED SEEDS                                                -->
  <!-- ============================================================ -->
  <deployedseed>12345</deployedseed>
  <deployedseed>23456</deployedseed>
  <deployedseed>34567</deployedseed>
  <deployedseed>45678</deployedseed>
  <deployedseed>56789</deployedseed>
  <deployedseed>67890</deployedseed>

  <!-- ============================================================ -->
  <!-- QUESTION TESTS (6 total)                                     -->
  <!-- ============================================================ -->

  <!-- Test 1: All correct answers -->
  <qtest>
    <testcase>1</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>ta6</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>ta8</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt6-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt8-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 2: Wrong MCQ (ans1=2) -->
  <qtest>
    <testcase>2</testcase>
    <testinput>
      <name>ans1</name>
      <value>2</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>ta6</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>ta8</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.1000000</expectedpenalty>
      <expectedanswernote>prt1-0-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt6-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt8-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 3: Wrong i(0+) (ans2=1) -->
  <qtest>
    <testcase>3</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>1</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>ta6</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>ta8</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.1000000</expectedpenalty>
      <expectedanswernote>prt2-0-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt6-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt8-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 4: Wrong vC(0+) (ans3=0) -->
  <qtest>
    <testcase>4</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>0</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>ta6</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>ta8</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.1000000</expectedpenalty>
      <expectedanswernote>prt3-0-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt6-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt8-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 5: Wrong regime (always wrong regardless of set) -->
  <qtest>
    <testcase>5</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>mod(ta_regime,3)+1</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>ta8</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.1000000</expectedpenalty>
      <expectedanswernote>prt6-0-F</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt8-0-T</expectedanswernote>
    </expected>
  </qtest>

  <!-- Test 6: Wrong i(t) (ans8=0) -->
  <qtest>
    <testcase>6</testcase>
    <testinput>
      <name>ans1</name>
      <value>ta1</value>
    </testinput>
    <testinput>
      <name>ans2</name>
      <value>ta2</value>
    </testinput>
    <testinput>
      <name>ans3</name>
      <value>ta3</value>
    </testinput>
    <testinput>
      <name>ans4</name>
      <value>ta4</value>
    </testinput>
    <testinput>
      <name>ans5</name>
      <value>ta5</value>
    </testinput>
    <testinput>
      <name>ans6</name>
      <value>ta6</value>
    </testinput>
    <testinput>
      <name>ans7</name>
      <value>ta7</value>
    </testinput>
    <testinput>
      <name>ans8</name>
      <value>0</value>
    </testinput>
    <expected>
      <name>prt1</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt1-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt2</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt2-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt3</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt3-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt4</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt4-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt5</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt5-1-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt6</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt6-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt7</name>
      <expectedscore>1.0000000</expectedscore>
      <expectedpenalty>0.0000000</expectedpenalty>
      <expectedanswernote>prt7-0-T</expectedanswernote>
    </expected>
    <expected>
      <name>prt8</name>
      <expectedscore>0.0000000</expectedscore>
      <expectedpenalty>0.1000000</expectedpenalty>
      <expectedanswernote>prt8-0-F</expectedanswernote>
    </expected>
  </qtest>

  <!-- ============================================================ -->
  <!-- HINTS (3, progressive)                                        -->
  <!-- ============================================================ -->

  <hint format="html">
    <text><![CDATA[<p><strong>Hint 1 (Intuition):</strong> In DC steady state, think about what happens to each component after an infinite amount of time. Which component prevents DC current from flowing?</p>]]></text>
  </hint>
  <hint format="html">
    <text><![CDATA[<p><strong>Hint 2 (Formulas):</strong> Capacitor = open circuit at DC (blocks DC current). Inductor = short circuit at DC (zero voltage drop). For a series RLC circuit: \(\alpha = \frac{R}{2L}\), \(\omega_0 = \frac{1}{\sqrt{LC}}\).</p>]]></text>
  </hint>
  <hint format="html">
    <text><![CDATA[<p><strong>Hint 3 (Worked Step):</strong> With no current flowing before the switch opens: \(i(0^-) = 0\;\text{A}\) and \(v_C(0^-) = V_s\). By continuity these values hold at \(t = 0^+\). Then from KVL after the switch opens:</p><p>\[\frac{di}{dt}(0^+) = \frac{1}{L}\left(-R \cdot i(0^+) - v_C(0^+)\right) = \frac{-V_s}{L}\]</p>]]></text>
  </hint>

</question>

</quiz>
